rule_files:
    - alerts.rules_test.yml

evaluation_interval: 1m

tests:
    # Test 1.
    - interval: 1m
      input_series:
          - series: 'probe_success{instance="oncall-web:8080"}'
            values: '0 0 0 0 0 0 0 0 0 0 0 0 0 0 0'
      alert_rule_test:
          - eval_time: 10m
            alertname: OncallDown
            exp_alerts:
                - exp_labels:
                      severity: critical
                      instance: oncall-web:8080
                  exp_annotations:
                      description: "Oncall Web is down for more than 1 minute."

    # Test 2.
    - interval: 1m
      input_series:
        - series: 'duty{role="primary"}'
          values: '0 0 0 0 0 0 0 0 0 0 0 0 0 0 0'
      alert_rule_test:
        - eval_time: 10m
          alertname: NoDutyPrimary
          exp_alerts:
            - exp_labels:
                severity: critical
                role: primary
              exp_annotations:
                description: "No duty primary in team."

    # Test 3.
    - interval: 1m
      input_series:
        - series: 'node_filesystem_avail{job="node-exporter"}'
          values: '1 5 10 15 19'
        - series: 'node_filesystem_size{job="node-exporter"}'
          values: '100 100 100 100 100'
      alert_rule_test:
        - eval_time: 5m
          alertname: DiskSpaceUsed
          exp_alerts:
            - exp_labels:
                severity: warning
                job: node-exporter
              exp_annotations:
                description: "Disk space used 80%"

    # Test 4.
    - interval: 1m
      input_series:
        - series: 'duty{role="primary"}'
          values: '0 0 1'
        - series: 'duty{role="secondary"}'
          values: '0 1 0'
      alert_rule_test:
        - eval_time: 5m
          alertname: NoDutyWeekSecondary
          exp_alerts:
            - exp_labels:
                severity: warning
                role: secondary
              exp_annotations:
                description: "Duty Week is not primary or secondary in team."
